---
title        : Template Chapter
description  : This is a template chapter. 

--- type:NormalExercise lang:python xp: skills: key:4ace9cb3ae
## Template exercise 
Thisis a template exercise.

*** =instructions

*** =hint

*** =pre_exercise_code{python}

*** =sample_code{python}

*** =solution{python}

*** =sct{python}
